rule_files:
  - alerts.yml

tests:
  - interval: 5m
    input_series:
      - series: 'ethereum_steth_network_errors_total{}'
        values: '100 110 120 125 130 165'
    alert_rule_test:
      - alertname: TooMuchNetworkErrors
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: 'Eth-steth has detected 25 network errors in the last 15 minutes'
              description: 'The ethereum_steth_network_errors_total 25 in the last 15 minutes.'

      - alertname: TooMuchNetworkErrors
        eval_time: 20m
        exp_alerts: []

      - alertname: TooMuchNetworkErrors
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: 'Eth-steth has detected 45 network errors in the last 15 minutes'
              description: 'The ethereum_steth_network_errors_total 45 in the last 15 minutes.'
